# Activity: Removing Data

## Introduction

In this activity you will practice using Pandas functionality to check for and remove any unwanted data from a dataset.
This activity will cover the following topics:
- Removing columns from a DataFrame
- Removing rows from a DataFrame
- Removing rows based on a condition
- Checking for duplicate data


#### Question 1

Create a `DataFrame` called `df` from the given CSV file `exotic_plants_data.csv`, then drop the column `Type` and assign the result to a new `DataFrame` called `df_no_type`.


```python
ul="https://raw.githubusercontent.com/ddaeducation/datanalytics/main/exotic_plants_data.csv"
```

```python
data=pd.read_csv(url)
data.head(10)
```

Output:
```
Title  Year Released  Rating  \
0                           The Shawshank Redemption           1994     9.3   
1                                      The Godfather           1972     9.2   
2                                    The Dark Knight           2008     9.0   
3                                       Pulp Fiction           1994     8.9   
4                                   Schindler's List           1993     8.9   
5                                         Fight Club           1999     8.8   
6                                       Forrest Gump           1994     8.8   
7                                          Inception           2010     8.8   
8                                         The Matrix           1999     8.7   
9  The Lord of the Rings: The Fellowship of the Ring           2001     8.8   

   Box Office ($M)  
0             58.3  
1            246.1  
2           1005.0  
3            213.9  
4            321.3  
5            100.9  
6            678.2  
7        
```

<!-- grade_id: cell-99d36e8c8c66ea51, points: 0, grade: false, solution: true -->
```python
import pandas as pd


# Your code here

# your code here


```

<!-- grade_id: cell-84571c0862a237ac, points: 1, grade: true, solution: false -->
```python
# Question 1 Grading Checks

assert isinstance(df, pd.DataFrame), 'Have you created a DataFrame named df?'
assert isinstance(df_no_type, pd.DataFrame), 'Have you created a DataFrame named df_no_type?'

```

#### Question 2

Check the `df` `DataFrame` for any duplicate rows and assign the result to a new `DataFrame` called `df_duplicates`.


<!-- grade_id: cell-ee673ddb1d8d4343, points: 0, grade: false, solution: true -->
```python
# Your code here

# your code here


```

<!-- grade_id: cell-496fa33508cec3f2, points: 1, grade: true, solution: false -->
```python
# Question 2 Grading Checks

assert isinstance(df_duplicates, pd.DataFrame), 'Have you created a DataFrame named df_duplicates?'

```

#### Question 3

Check the `df` `DataFrame` for any duplicate rows based on the `Plant Name` and `Type` columns and assign the result to a new `DataFrame` called `df_plant_type_duplicates`.


<!-- grade_id: cell-6f7859a5535a9c9b, points: 0, grade: false, solution: true -->
```python
# Your code here

# your code here


```

<!-- grade_id: cell-67e41e79158b9f13, points: 1, grade: true, solution: false -->
```python
# Question 3 Grading Checks

assert isinstance(df_plant_type_duplicates, pd.DataFrame), 'Have you created a DataFrame named df_duplicates?'

```

#### Question 4

Create a mask called `clean_mask` that will clean up any duplicates in the `df` DataFrame that have the same `Plant Name` and `Origin` and only keep the most up-to-date duplicate entry.


<!-- grade_id: cell-722a62389afc976a, points: 0, grade: false, solution: true -->
```python
# Your code here

# your code here


```

<!-- grade_id: cell-a3bd7512e6b2d51c, points: 1, grade: true, solution: false -->
```python
# Question 4 Grading Checks

assert isinstance(clean_mask, pd.Series), 'Have you created a Series named clean_mask?'

```